## Εργασία 2 ##

Καλωσήρθατε στην δεύτερη σας εργασία. Η εργασία αυτή έχει σκοπό να σας βοηθήσει να εμπεδώσετε τα δενδρικά μοντέλα και την αξιολόγηση μοντέλων.

Στην εργασία αυτή θα πρέπει να συμπληρώσετε κώδικα Python 3 στα σημεία που αναφέρουν # YOUR CODE HERE. Μην τροποποιείτε τον κώδικα που βρίσκεται εκτός αυτών των περιοχών.

Πρωτού παραδόσετε την εργασία σας σιγουρευτείτε ότι ο κώδικας σε όλα τα κελιά τρέχει σωστά. Για το σκοπό αυτό επιλέξτε από το μενού **Χρόνος εκτέλεσης (runtime) -> Επανεκίνηση περιόδου λειτουργίας και εκτέλεση όλων**.

Συμπληρώστε το όνομα (NAME) και το AEM σας παρακάτω:

```python
NAME = "Βλαχάκης Θεόδωρος"
AEM = "10476"
```

**1 (0,5 μονάδες):** Διαβάστε το διαθέσιμο από την sklearn σύνολο δεδομένων diabetes που αφορά την πρόβλεψη της πορείας που θα έχει ένας ασθενής με διαβήτη ένα χρόνο μετά (εργασία παλινδρόμησης) και αποθηκεύστε το σε μια μεταβλητή με το όνομα diabetes.

<!-- grade_id: cell-945e6199220ebe51, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

diabetes = datasets.load_diabetes()
```

<!-- grade_id: cell-5ad298e30b073224, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθής ανάγνωσης του συνόλου δεδομένων"""
assert diabetes.feature_names == ['age', 'sex', 'bmi', 'bp', 's1', 's2', 's3', 's4', 's5', 's6']
```

**2 (0,5 μονάδες):** Χωρίστε το σύνολο που διαβάσατε σε υποσύνολο ελέγχου 20% και υποσύνολο εκπαίδευσης 80% με την μέθοδο train_test_split και σπόρο τυχαιότητας 42. Aποθηκεύστε το σύνολο εκπαίδευσης σε μεταβλητές X_train, y_train και το σύνολο ελέγχου σε μεταβλητές X_test, y_test. (0.5 μονάδες)

<!-- grade_id: cell-007e5ef64c2edaaf, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split

X = diabetes.data
y = diabetes.target
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)
```

<!-- grade_id: cell-33de7b4d5141b07e, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == 0.02736
assert round(X_test[0][8], 5) == 0.03243
```

**3 (1 μονάδα):** Χρησιμοποιώντας τα δεδομένα εκπαίδευσης, εκπαιδεύστε δύο δένδρα παλινδρόμησης με κριτήρια διαχωρισμού τα squared_eror και poisson ([δείτε ενότητα 1.10.7.2](https://scikit-learn.org/stable/modules/tree.html#tree)) αντίστοιχα, θέτοντας τον σπόρο τυχαιότητας στην τιμή 42 και το μέγιστο βάθος στην τιμή 3. Υπολογίστε τη μετρική r2 με βάση τις προβλέψεις τους στα δεδομένα ελέγχου, τις οποίες αποθηκεύστε στις μεταβλητές y_pred1 (squared_error) και y_pred2 (poisson) αντίστοιχα.
Αποθηκεύστε τις μετρικές στις μεταβλητές r2_squared_error και r2_poisson αντίστοιχα. (1 μονάδα)

<!-- grade_id: cell-b1390f68c6c3dc60, points: 0, grade: false, solution: true -->
```python
from sklearn.tree import DecisionTreeRegressor
from sklearn.metrics import r2_score

model_tree_regressor_squared_error = DecisionTreeRegressor(criterion='squared_error', max_depth=3, random_state=42)
model_tree_regressor_poisson = DecisionTreeRegressor(criterion='poisson', max_depth=3, random_state=42)

# Fit the models to the training set and make predictions using them

model_tree_regressor_squared_error.fit(X_train, y_train)
y_pred1 = model_tree_regressor_squared_error.predict(X_test)
r2_squared_error = r2_score(y_test, y_pred1)

model_tree_regressor_poisson.fit(X_train, y_train)
y_pred2 = model_tree_regressor_poisson.predict(X_test)
r2_poisson = r2_score(y_test, y_pred2)
```

<!-- grade_id: cell-3f2d723addec3c0f, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθού υπολογισμού των μετρικών"""
assert round(r2_squared_error, 3) == 0.329
assert round(r2_poisson, 3) == 0.391
assert len(y_pred1) == len(y_pred2)
```

**4 (0,5 μονάδες):** Διαχωρίστε τα δεδομένα diabetes σε υποσύνολα εκπαίδευσης 60% (X_train, y_train), επικύρωσης 15% (X_val, y_val) και ελέγχου 25% (X_test, y_test) με σπόρο τυχαιότητας 42. Ξεκινήστε δημιουργώντας πρώτα το τελικό σύνολο ελέγχου, και έπειτα χωρίζοντας το υπόλοιπο σε σύνολο εκπαίδευσης και επικύρωσης.

<!-- grade_id: cell-db6d2d4c83a94e4b, points: 0, grade: false, solution: true -->
```python
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.25, random_state=42)
X_train, X_val, y_train, y_val = train_test_split(X_train, y_train, test_size=0.2, random_state=42)
```

<!-- grade_id: cell-262a292a0b3f72bb, points: 0.5, grade: true, solution: false -->
```python
"""Τεστ ορθού διαχωρισμού του συνόλου δεδομένων"""
assert round(X_train[0][8], 5) == -0.08024
assert round(X_val[0][8], 5) == 0.09924
assert round(X_test[0][8], 5) == 0.03243
```

**5 (2 μονάδες):** Υλοποιήστε συνάρτηση, η οποία δοθέντος ενός δενδρικού μοντέλου παλινδρόμησης, διερευνά όλες τις ακέραιες τιμές της μεταβλητής min_samples_leaf από ένα κάτω όριο μέχρι ένα άνω όριο και επιστρέφει την βέλτιστη με βάση το mean squared error στο παραπάνω σύνολο επικύρωσης έπειτα από εκπαίδευση στο παραπάνω σύνολο εκπαίδευσης. Επιπλέον επιστρέφει και το mean squared error (στρογγυλοποιημένο σε ακέραιο) που πετυχαίνει στο σύνολο ελέγχου το μοντέλο με την βέλτιστη επίδοση στο σύνολο επικύρωσης.

<!-- grade_id: cell-db6235d877cf3abe, points: 0, grade: false, solution: true -->
```python
from sklearn.metrics import mean_squared_error

def best_min_samples_leaf(model, min, max):
    """
    The best_min_samples_leaf function explores all possible integer values of the 'min_samples_leaf' parameter within
    a specified range (from min to max), in order to find the optimal value for a DecisionTreeRegressor model.
    The function best_min_samples_leaf fits a DecisionTreeRegressor model, that is given as input, to the training set
    (X_train, y_train) and searches for the optimal value of the 'min_samples_leaf' parameter based on the Mean Squared
    Error (MSE) calculated on the validation set (X_val, y_val).
    Once the optimal 'min_samples_leaf' value is found, the function sets this value to the DecisionTreeRegressor model's
    parameter and retrains the model on the training set.
    The function best_min_samples_leaf returns the optimal value of the 'min_samples_leaf' parameter as well as the MSE
    that the model with this optimal performance on the validation set achieves on the test set (X_test, y_test).

    Inputs:
    - model: A scikit-learn DecisionTreeRegressor instance.
    - min: The minimum integer value to consider for the 'min_samples_leaf' parameter.
    - max: The maximum integer value to consider for the 'min_samples_leaf' parameter.

    Outputs:
    - min_samples_leaf_opt: The optimal value of the 'best_min_samples_leaf' parameter for the DecisionTreeRegressor model.
    - mean_squared_error_opt: The MSE metric value (rounded to integer) that the model with the optimal performance on the
                              validation seτ achieves on the test set.
    """

    #params = {"min_samples_leaf": range(min, max+1)}

    min_mse_on_val_set = float('inf')
    min_samples_leaf_opt = min

    # Find the optimal value of the 'min_samples_leaf' parameter for the model based on MSE calculated on the validation set
    for i in range(min, max+1):
      model.set_params(min_samples_leaf = i)
      model.fit(X_train, y_train)

      temp_mse = mean_squared_error(y_val, model.predict(X_val))

      # If temp_mse is equal to the current minimum MSE (min_mse_on_val_set, which is defined using a previous model) then we
      # don't update the optimal value of 'min_samples_leaf' parameter for the model, since we prefer the smaller 'min_samples_leaf'
      # value in case of a tie. So in the condition below we search using strict inequality.
      if(temp_mse < min_mse_on_val_set):
        min_samples_leaf_opt = i
        min_mse_on_val_set = temp_mse

    # Once the optimal 'min_samples_leaf' value is found, redefine this optimal model and retrain it on the training set.
    # Then calculate the MSE that the optimal model achieves on the test set
    model.set_params(min_samples_leaf = min_samples_leaf_opt)
    model.fit(X_train, y_train)
    mean_squared_error_opt = round(mean_squared_error(y_test, model.predict(X_test)))

    return min_samples_leaf_opt, mean_squared_error_opt
```

<!-- grade_id: cell-4c4b39c4be673d88, points: 2, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_depth=3, random_state=42)
assert best_min_samples_leaf(model, 1, 8) == (7, 3358)
assert best_min_samples_leaf(model, 5, 15) == (11, 3271)
```

**6 (1,5 μονάδα):** Υλοποιήστε συνάρτηση η οποία δοθέντος ενός μοντέλου, ενός συνόλου δεδομένων και μίας μετρικής, αξιολογεί το μοντέλο με τη μέθοδο της σταυρωτής επικύρωσης "άφησε ένα εκτός".

<!-- grade_id: cell-83aa7655f3f1ecdd, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import LeaveOneOut
import numpy as np

def leave_one_out(model, X, y, metric):
    """
    The leave_one_out function evaluates a given model on a dataset and a specified metric, applying the Leave One Out Cross-Validation method.

    Inputs:
    - model: The given model to be evaluated.
    - X: A 2D array which represents the validation/testing dataset (examples).
    - y: A 1D array which represents the target (output) values corresponding to the input data (X) in the dataset.
    - metric: The metric with which the model is being evaluated.

    Outputs:
    - metric_value: A float number which represents the model's performance score, according to the specified (given) evaluation metric.
    """

    # Define a "Leave One Out" object to apply this cross validation method. The Leave One Out Cross-Validation method splits the given
    # dataset into m subsets of data (where m is the actual size of the given dataset) and executes a repeatable process. In each iteration
    # of this process, only one subset (i.e example) is considered as validation/testing set and the union of the rest is considered as
    # the training set. So, the model is trained on this training set and evaluated on the one example left out (current test set). This
    # process is being repeated for all the subsets of data.
    loo = LeaveOneOut()

    accuracies = []
    for train_index, test_index in loo.split(X, y):
        X_train, X_test = X[train_index], X[test_index]
        y_train, y_test = y[train_index], y[test_index]

        model.fit(X_train, y_train)
        y_pred = model.predict(X_test)

        accuracies.append(metric(y_test, y_pred))

    # Compute the model's performance score accorind to the specified metric, by averaging the accuracies
    metric_value = np.average(accuracies)
    return metric_value
```

<!-- grade_id: cell-06bac2fcee9a42de, points: 1.5, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_leaf_nodes=6, min_samples_leaf=2, random_state=42)
assert round(leave_one_out(model, diabetes.data, diabetes.target, mean_squared_error), 0) == 3354
```

**7 (2 μονάδες):** Χρησιμοποιώντας τις κλάσεις GridSearchCV, KFold και τη συνάρτηση cross_val_score, υλοποιήστε συνάρτηση που δέχεται δενδρικό μοντέλο, σύνολο δεδομένων, μετρική αξιολόγησης και ελάχιστη (min) και μέγιστη (max) τιμή βάθους για το δένδρο και επιστρέφει εκτίμηση της απόδοσης του μοντέλου αυτού, έπειτα από επιλογή της καλύτερης δυνατής τιμής για την παράμετρο του βάθους από min έως και max, με βάση τη δοθείσα μετρική και τη μέθοδο της φωλιασμένης k-πλης σταυρωτής επικύρωσης. Χρησιμοποιήστε 10 fold εξωτερικά, 9 fold εσωτερικά, τυχαίο ανακάτεμα δεδομένων, και σπόρο τυχαιότητας 25.

<!-- grade_id: cell-0859aa898e82771f, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import GridSearchCV, KFold, cross_val_score
import numpy as np

def estimation(model, X, y, metric, min, max):
    """
    The "estimation" function evaluates a given tree model on a given dataset, according to a specified metric, first
    exploring all possible integer values of the 'max_depth' parameter of the model within a specified range (from min
    to max), in order to find the optimal value of this parameter for the model and then estimating the performance of
    this optimal tree model on the given dataset. So, the function executes both model selection and estimation of its
    performance.
    Both model selection and performance estimation is being executed according to the specified metric.
    The function executes the entire evaluation (selection and model's performance estimation) of the model using the
    Nested K-Fold Cross-Validation method, with 10 outer folds and 9 inner folds. For each outer split, the function
    performs model selection by testing all integer values of the 'max_depth' parameter of the tree model within a
    specified range (from min to max) using the inner Cross-Validation. Once the optimal 'max_depth' value is found,
    the optimal model is being selected and it is then evaluated on the corresponding outer fold of the dataset.

    Inputs:
    - model: The tree model to be evaluated.
    - X: A 2D array which represents the dataset on which the model will be trained.
    - y: A 1D array which represents the target (output) values corresponding to the input data in the training set X.
    - metric: The metric with which the model is being evaluated.
    - min: The minimum integer value to consider for the 'max_depth' parameter of the tree model.
    - max: The maximum integer value to consider for the 'max_depth' parameter of the tree model.

    Outputs:
    - model_evaluation_score: A float number which represents estimation for the model's (the model selected as optimal according to the
                              specified metric, by choosing the corresponding value of 'max_depth' parameter) performance score, according
                              to the specified evaluation metric
    """

    params = {'max_depth': range(min, max+1)}

    outer_cv = KFold(n_splits=10, shuffle=True, random_state=25)
    inner_cv = KFold(n_splits=9, shuffle=True, random_state=25)

    gs_model_selection = GridSearchCV(estimator=model, param_grid=params, cv=inner_cv)

    performance_scores = cross_val_score(gs_model_selection, X, y, scoring=metric, cv=outer_cv)

    # Now average the performance scores of the estimator for each iteration of outer cross validation in order to get an estimation for the
    # pefromance score of the selected model.
    # Use abs in order to return the absolute value of the estimated score.
    performance_score = abs(np.average(performance_scores))

    return performance_score
```

<!-- grade_id: cell-29fa8f9d02b1331e, points: 2, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(random_state=42)
assert round(estimation(model, diabetes.data, diabetes.target, 'neg_mean_squared_error', 1, 10), 2) == 3993.82
```

**8 (1 μονάδα):** Υλοποιήστε συνάρτηση που δέχεται εκπαιδευμένο δενδρικό μοντέλο και πίνακα συμβολοσειρών με τα ονόματα των μεταβλητών εισόδου και επιστρέφει τη σημαντικότερη μεταβλητή εισόδου.

<!-- grade_id: cell-df6dc2f47f9d098b, points: 0, grade: false, solution: true -->
```python
def best_feature(model, feature_names):
    """
    Given a trained tree model and an array including the strings with the names of the model's input variables, the best_feature
    function estimates which is the most important input variable in this model and returns the string with the name of that variable.
    In terms of the importance of an input variable, the most important input variable of the tree model is the one that achieves
    the largest relative reduction (i.e., percentage reduction, by expanding the node of the tree based on this variable, in relation
    to the percentage of examples that relate to this node) of the eterogeneity of training examples from expanding the specific node
    based on this variable.

    Inputs:
    - model: The trained tree model whose most important input variable is wanted.
    - feature_names: List of strings with the names of the model's input variables.

    Ouputs:
    - most_important_var: The most important input variable.
    """
    # Since the tree model is already trained in some training set, we can call its attribute feature_importances_ which is a 1D array
    # that stores an importance score, in terms of importance of an input variable, for each input variable in the training set.
    most_important_var_idx = np.argmax(model.feature_importances_)
    most_important_var = feature_names[most_important_var_idx]

    return most_important_var

```

<!-- grade_id: cell-50285e47d84e6828, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_depth=3, random_state=42)
model.fit(diabetes.data, diabetes.target)
assert best_feature(model, diabetes.feature_names) == 's5'
```

**9 (1 μονάδα):** Υλοποιήστε συνάρτηση που δέχεται εκπαιδευμένο δενδρικό μοντέλο και πίνακα συμβολοσειρών με τα ονόματα των μεταβλητών και επιστρέφει την μεταβλητή η οποία ελέγχεται στο δένδρο στο πρώτο παιδί δεξιά. Χρησιμοποιήστε την ιδιότητα [tree_](https://scikit-learn.org/stable/auto_examples/tree/plot_unveil_tree_structure.html#sphx-glr-auto-examples-tree-plot-unveil-tree-structure-py).

<!-- grade_id: cell-ead77d5ea43abcf1, points: 0, grade: false, solution: true -->
```python
def right_child(model, feature_names):
    """
    Given a trained tree model and an array including the strings with the names of the model's input variables, the right_child
    function looks for the input variable that is being examined in the tree model in the first child on the right.

    Inputs:
    - model: The trained tree model.
    - feature_names: List of strings with the names of the model's input variables.

    Ouputs:
    - first_child_on_the_right_var: The name of the variable being checked in the tree in the first chlid on the right.
    """

    # At first, check that the right child is root of a sub-tree expanding from this according to a feature.
    if(model.max_depth == 1):
      return "There is no feature check on the first right child of the tree since it is a leaf"
    if(model.max_depth == 0):
      return "There is not even a first child on the right in the tree. The root of the tree is a leaf."

    # Firstly, find the id of the node which is the root of the sub-tree expanding from the first right child. In other words find
    # the id of the node which is the right child of the root of the tree model.
    first_right_child_id = model.tree_.children_right[0]

    # Now find the feature used for splitting the node with the id found.
    first_child_on_the_right_var_idx = model.tree_.feature[first_right_child_id]

    first_child_on_the_right_var = feature_names[first_child_on_the_right_var_idx]

    return first_child_on_the_right_var
```

<!-- grade_id: cell-c59ba7d856846a2c, points: 1, grade: true, solution: false -->
```python
"""Τεστ ορθής υλοποίησης συνάρτησης"""
model = DecisionTreeRegressor(max_depth=2, random_state=42)
model.fit(diabetes.data, diabetes.target)
assert right_child(model, diabetes.feature_names) == 'bmi'
```